Demo Manually graded answer

Please state Newton's three law of motion:

<!-- grade_id: newton_law_motion, points: 1, grade: true, solution: true -->
```python

```

Demo Autograded answer

Please write the Linear Momentum formula in the below cell:

<!-- grade_id: momentum_formula, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
p = mv
### END SOLUTION
```

Demo Autograder tests

<!-- grade_id: cell-d4a12ebfd43a43ba, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
a = 5
assert a == 5
### END HIDDEN TESTS
```

Demo Read-only

```python
a = 1 + 1
assert a == 2
```

Manually graded task

<!-- grade_id: cell-f007c5d5b4e88524, points: 1, grade: true, solution: true -->
Describe the task here!