```python
from numpy.testing import assert_allclose
```

Excercise 1

a) Create a variable t with a value of 5

<!-- grade_id: variable, points: 0, grade: false, solution: true -->
```python
# enter you answer below this line
### BEGIN SOLUTION
t = 5
### END SOLUTION
```

b) The code below should run without error

<!-- grade_id: cell-d4a12ebfd43a43ba, points: 1, grade: true, solution: false -->
```python
s = 10/t
assert_allclose(s, 2.0, rtol=1e-2)
```